<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.4.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac10_4_3">
<div class="ac_question col-md-12" id="foppff_1_ac10_4_3_question"><div class="para" id="p-2513">
<dfn class="terminology">3.</dfn> Below, we have provided a list of lists named <code class="code-inline tex2jax_ignore">L</code>. Use nested iteration to save every string containing “b” into a new list named <code class="code-inline tex2jax_ignore">b_strings</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac10_4_3_editor" data-question_label="" data-codelens="true">L = [['apples', 'bananas', 'oranges', 'blueberries', 'lemons'], ['carrots', 'peas', 'cucumbers', 'green beans'], ['root beer', 'smoothies', 'cranberry juice']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(b_strings, ['bananas', 'blueberries', 'cucumbers', 'green beans', 'root beer', 'cranberry juice'], "Testing that b_strings was created correctly.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_nested-iteration.html#ac10_4_3" class="internal">in-context</a></span>
</body>
</html>
