<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h4 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.1.2</span><span class="period">.</span>
</h4>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac10_1_6">
<div class="ac_question col-md-12" id="foppff_1_ac10_1_6_question"><div class="para" id="p-2474">
<dfn class="terminology">2.</dfn> Using indexing, retrieve the string ‘willow’ from the list and assign that to the variable <code class="code-inline tex2jax_ignore">plant</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac10_1_6_editor" data-question_label="" data-codelens="true">data = ['bagel', 'cream cheese', 'breakfast', 'grits', 'eggs', 'bacon', [34, 9, 73, []], [['willow', 'birch', 'elm'], 'apple', 'peach', 'cherry']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(plant, 'willow', "Testing that plant has the correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_introduction-nested-data-and-nested-iteration.html#ac10_1_6" class="internal">in-context</a></span>
</body>
</html>
