<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h4 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.1.1</span><span class="period">.</span>
</h4>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac10_1_5">
<div class="ac_question col-md-12" id="foppff_1_ac10_1_5_question"><div class="para" id="p-2473">
<dfn class="terminology">1.</dfn> Below, we have provided a list of lists. Use indexing to assign the element ‘horse’ to the variable name <code class="code-inline tex2jax_ignore">idx1</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac10_1_5_editor" data-question_label="" data-codelens="true">animals = [['cat', 'dog', 'mouse'], ['horse', 'cow', 'goat'], ['cheetah', 'giraffe', 'rhino']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(idx1, 'horse', "Testing that idx1 was assigned correctly.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_introduction-nested-data-and-nested-iteration.html#ac10_1_5" class="internal">in-context</a></span>
</body>
</html>
