<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">11.3.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac11_3_2">
<div class="ac_question col-md-12" id="foppff_1_ac11_3_2_question"><div class="para logical" id="p-2559"><ol class="decimal"><li id="li-294"><div class="para" id="p-2560">Find the number of lines in the file, <code class="code-inline tex2jax_ignore">travel_plans2.txt</code>, and assign it to the variable <code class="code-inline tex2jax_ignore">num_lines</code>.</div></li></ol></div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac11_3_2_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testTwo(self):
      self.assertEqual(num_lines, 11, "Testing that num_lines is assigned to correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="files_alternative-file-reading-methods.html#ac11_3_2" class="internal">in-context</a></span>
</body>
</html>
