<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">11.3.5</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac11_3_3">
<div class="ac_question col-md-12" id="foppff_1_ac11_3_3_question"><div class="para logical" id="p-2561"><ol class="decimal"><li id="li-295"><div class="para" id="p-2562">Create a string called <code class="code-inline tex2jax_ignore">first_forty</code> that is comprised of the first 40 characters of <code class="code-inline tex2jax_ignore">emotion_words2.txt</code>.</div></li></ol></div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac11_3_3_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(first_forty, 'Sad upset blue down melancholy somber bi', "Testing that first_forty was created correctly.")
myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="files_alternative-file-reading-methods.html#ac11_3_3" class="internal">in-context</a></span>
</body>
</html>
