<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">11.14.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac11_14_2">
<div class="ac_question col-md-12" id="foppff_1_ac11_14_2_question"><div class="para" id="p-2654">We have provided a file called <code class="code-inline tex2jax_ignore">emotion_words.txt</code> that contains lines of words that describe emotions. Find the total number of words in the file and assign this value to the variable <code class="code-inline tex2jax_ignore">num_words</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac11_14_2_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(num_words, 48, "Testing that num_words was assigned to the correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="files_chapter-assessment.html#ac11_14_2" class="internal">in-context</a></span>
</body>
</html>
