<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">11.14.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac11_14_4">
<div class="ac_question col-md-12" id="foppff_1_ac11_14_4_question"><div class="para" id="p-2656">Assign the first 30 characters of <code class="code-inline tex2jax_ignore">school_prompt.txt</code> as a string to the variable <code class="code-inline tex2jax_ignore">beginning_chars</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac11_14_4_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(len(beginning_chars), 30, "Testing that beginning_chars has the correct length.")
      self.assertEqual(beginning_chars, "Writing essays for school can ", "Testing that beginning_chars has the correct string.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="files_chapter-assessment.html#ac11_14_4" class="internal">in-context</a></span>
</body>
</html>
