<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">12.3.8</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac12_04_8">
<div class="ac_question col-md-12" id="foppff_1_ac12_04_8_question"><div class="para" id="p-2681">Provided is a buggy for loop that tries to accumulate some values out of some dictionaries. Insert a try/except so that the code passes. If the key is not there, initialize it in the dictionary and set the value to zero.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac12_04_8_editor" data-question_label="" data-codelens="true">di = [{"Puppies": 17, 'Kittens': 9, "Birds": 23, 'Fish': 90, "Hamsters": 49}, {"Puppies": 23, "Birds": 29, "Fish": 20, "Mice": 20, "Snakes": 7}, {"Fish": 203, "Hamsters": 93, "Snakes": 25, "Kittens": 89}, {"Birds": 20, "Puppies": 90, "Snakes": 21, "Fish": 10, "Kittens": 67}]
total = 0
for diction in di:
    total = total + diction['Puppies']

print("Total number of puppies:", total)


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      accum = 0
      for diction in di:
           if 'Puppies' in diction:
               accum += 1
      self.assertEqual(accum, 4, "Testing that every dictionary in di has the key 'Puppies'.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="exceptions_chapter-assessment.html#ac12_04_8" class="internal">in-context</a></span>
</body>
</html>
