<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">12.3.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac12_04_4">
<div class="ac_question col-md-12" id="foppff_1_ac12_04_4_question"><div class="para" id="p-2677">Provided is buggy code, insert a try/except so that the code passes.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac12_04_4_editor" data-question_label="" data-codelens="true">lst = [2,4,10,42,12,0,4,7,21,4,83,8,5,6,8,234,5,6,523,42,34,0,234,1,435,465,56,7,3,43,23]

lst_three = []

for num in lst:
    if 3 % num == 0:
        lst_three.append(num)


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(lst_three, [1,3], "Testing that lst_three has the correct values.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="exceptions_chapter-assessment.html#ac12_04_4" class="internal">in-context</a></span>
</body>
</html>
