<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">12.3.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac12_04_3">
<div class="ac_question col-md-12" id="foppff_1_ac12_04_3_question"><div class="para" id="p-2676">The list, <code class="code-inline tex2jax_ignore">numb</code>, contains integers. Write code that populates the list <code class="code-inline tex2jax_ignore">remainder</code> with the remainder of 36 divided by each number in <code class="code-inline tex2jax_ignore">numb</code>. For example, the first element should be 0, because 36/6 has no remainder. If there is an error, have the string “Error” appear in the <code class="code-inline tex2jax_ignore">remainder</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac12_04_3_editor" data-question_label="" data-codelens="true">numb = [6, 0, 36, 8, 2, 36, 0, 12, 60, 0, 45, 0, 3, 23]

remainder = []

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOneA(self):
      self.assertEqual(remainder, [0, 'Error', 0, 4, 0, 0, 'Error', 0, 36, 'Error', 36, 'Error', 0, 13], "Testing that remainder is assigned to correct values.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="exceptions_chapter-assessment.html#ac12_04_3" class="internal">in-context</a></span>
</body>
</html>
