<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">13.13.1</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac13_13_1">
<div class="ac_question col-md-12" id="foppff_1_ac13_13_1_question"><div class="para" id="p-2800">Define a class called <code class="code-inline tex2jax_ignore">Bike</code> that accepts a string and a float as input, and assigns those inputs respectively to two instance variables, <code class="code-inline tex2jax_ignore">color</code> and <code class="code-inline tex2jax_ignore">price</code>. Assign to the variable <code class="code-inline tex2jax_ignore">testOne</code> an instance of <code class="code-inline tex2jax_ignore">Bike</code> whose color is <dfn class="terminology">blue</dfn> and whose price is <dfn class="terminology">89.99</dfn>. Assign to the variable <code class="code-inline tex2jax_ignore">testTwo</code> an instance of Bike whose color is <dfn class="terminology">purple</dfn> and whose price is <dfn class="terminology">25.0</dfn>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac13_13_1_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(testOne.color, "blue", "Testing that testOne has the correct color assigned.")
      self.assertEqual(testOne.price, 89.99, "Testing that testOne has the correct price assigned.")

   def testTwo(self):
      self.assertEqual(testTwo.color, "purple", "Testing that testTwo has the correct color assigned.")
      self.assertEqual(testTwo.price, 25.0, "Testing that testTwo has the correct color assigned.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="classes_chapter-assessment.html#ac13_13_1" class="internal">in-context</a></span>
</body>
</html>
