<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.7.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac17_6_4">
<div class="ac_question col-md-12" id="foppff_1_ac17_6_4_question"><div class="para" id="p-2529">Provided is a nested data structure. Follow the instructions in the comments below. Do not hard code.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac17_6_4_editor" data-question_label="" data-codelens="true">nested = {'data': ['finding', 23, ['exercises', 'hangout', 34]], 'window': ['part', 'whole', [], 'sum', ['math', 'calculus', 'algebra', 'geometry', 'statistics',['physics', 'chemistry', 'biology']]]}

# Check to see if the string 'data' is a key in nested, if it is, assign True to the variable data, otherwise assign False.

# Check to see if the integer 24 is in the value of the key data, if it is then assign to the variable twentyfour the value of True, otherwise False.

# Check to see that the string 'whole' is not in the value of the key window. If it's not, then assign to the variable whole the value of True, otherwise False.

# Check to see if the string 'physics' is a key in the dictionary nested. If it is, assign to the variable physics, the value of True, otherwise False.

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(data, True, "Testing that data has the correct value.")
   def testTwo(self):
      self.assertEqual(twentyfour, False, "Testing that twentyfour has the correct value.")
   def testThree(self):
      self.assertEqual(whole, False, "Testing that whole has the correct value.")
   def testFour(self):
      self.assertEqual(physics, False, "Testing that physics has the correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_chapter-assessment.html#ac17_6_4" class="internal">in-context</a></span>
</body>
</html>
