<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.7.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac17_6_2">
<div class="ac_question col-md-12" id="foppff_1_ac17_6_2_question"><div class="para" id="p-2527">Below, a list of lists is provided. Use in and not in tests to create variables with Boolean values. See comments for further instructions.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac17_6_2_editor" data-question_label="" data-codelens="true">lst = [['apple', 'orange', 'banana'], [5, 6, 7, 8, 9.9, 10], ['green', 'yellow', 'purple', 'red']]

#Test to see if 'yellow' is in the third list of lst. Save to variable ``yellow``


#Test to see if 4 is in the second list of lst. Save to variable ``four``


#Test to see if 'orange' is in the first element of lst. Save to variable ``orange``

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testTwo(self):
      self.assertEqual(yellow, True, "Testing that yellow is assigned to correct value")
   def testTwoB(self):
      self.assertEqual(four, False, "Testing that four is assigned to correct value")
   def testTwoC(self):
      self.assertEqual(orange, True, "Testing that orange is assigned to correct value")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_chapter-assessment.html#ac17_6_2" class="internal">in-context</a></span>
</body>
</html>
