<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.7.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac17_6_3">
<div class="ac_question col-md-12" id="foppff_1_ac17_6_3_question"><div class="para" id="p-2528">Below, we’ve provided a list of lists. Use in statements to create variables with Boolean values - see the ActiveCode window for further directions.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac17_6_3_editor" data-question_label="" data-codelens="true">L = [[5, 8, 7], ['hello', 'hi', 'hola'], [6.6, 1.54, 3.99], ['small', 'large']]

# Test if 'hola' is in the list L. Save to variable name test1

# Test if [5, 8, 7] is in the list L. Save to variable name test2

# Test if 6.6 is in the third element of list L. Save to variable name test3

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testA(self):
      self.assertEqual(test1, False, "Testing that test1 has the correct value.")
   def testB(self):
      self.assertEqual(test2, True, "Testing that test2 has the correct value.")
   def testC(self):
      self.assertEqual(test3, True, "Testing that test3 has the correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_chapter-assessment.html#ac17_6_3" class="internal">in-context</a></span>
</body>
</html>
