<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">10.7.8</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac17_6_8">
<div class="ac_question col-md-12" id="foppff_1_ac17_6_8_question"><div class="para" id="p-2533">Iterate through the contents of <code class="code-inline tex2jax_ignore">l_of_l</code> and assign the third element of sublist to a new list called <code class="code-inline tex2jax_ignore">third</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac17_6_8_editor" data-question_label="" data-codelens="true">l_of_l = [['purple', 'mauve', 'blue'], ['red', 'maroon', 'blood orange', 'crimson'], ['sea green', 'cornflower', 'lavender', 'indigo'], ['yellow', 'amarillo', 'mac n cheese', 'golden rod']]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(third, ['blue', 'blood orange', 'lavender', 'mac n cheese'], "Testing that third has the correct list assigned to it.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="nested-data_chapter-assessment.html#ac17_6_8" class="internal">in-context</a></span>
</body>
</html>
