<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.1</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_15_1">
<div class="ac_question col-md-12" id="foppff_1_ac4_15_1_question"><div class="para" id="p-1379">Write a function called <code class="code-inline tex2jax_ignore">int_return</code> that takes an integer as input and returns the same integer.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_15_1_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
    self.assertEqual(int_return(10), 10, "Testing that function int_return(10) returns 10")
myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_chapter-assessment.html#ac4_15_1" class="internal">in-context</a></span>
</body>
</html>
