<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.5</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_15_5">
<div class="ac_question col-md-12" id="foppff_1_ac4_15_5_question"><div class="para" id="p-1383">You will need to write two functions for this problem. The first function, <code class="code-inline tex2jax_ignore">divide</code> that takes in any number and returns that same number divided by 2. The second function called <code class="code-inline tex2jax_ignore">sum</code> should take any number, divide it by 2, and add 6. It should return this new number. You should call the <code class="code-inline tex2jax_ignore">divide</code> function within the <code class="code-inline tex2jax_ignore">sum</code> function. Do not worry about decimals.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_15_5_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testSixA(self):
    self.assertEqual(divide(4), 2, "Tests that divide(4) returns 2")
  def testSixB(self):
    self.assertEqual(sum(4), 8, "Tests that sum(4) returns 8")
    self.assertEqual(sum(2), 7, "Tests that sum(2) returns 7")
    self.assertEqual(sum(-6), 3, "Tests that sum(-6) returns 3")
    self.assertEqual(sum(0), 6, "Tests that sum(0) returns 6")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_chapter-assessment.html#ac4_15_5" class="internal">in-context</a></span>
</body>
</html>
