<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.10.1</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_10_1">
<div class="ac_question col-md-12" id="foppff_1_ac4_10_1_question"><div class="para" id="p-1289">
<dfn class="terminology">1.</dfn> Write two functions, one called <code class="code-inline tex2jax_ignore">addit</code> and one called <code class="code-inline tex2jax_ignore">mult</code>. <code class="code-inline tex2jax_ignore">
        addit</code> takes one number as an input and adds 5. <code class="code-inline tex2jax_ignore">mult</code> takes one number as an input, and multiplies that input by whatever is returned by <code class="code-inline tex2jax_ignore">addit</code>, and then returns the result.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_10_1_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
    self.assertEqual(mult(1), 6,"Testing the function mult with input 1 (should be 6)")
    self.assertEqual(mult(-2), -6, "Testing the function mult with input -2 (should be -6)")
    self.assertEqual(mult(0), 0, "Testing the function mult with input 0 (should be 0)")

  def testTwo(self):
    self.assertEqual(addit(1), 6, "Testing the function addit with input 1 (should be 6)")
    self.assertEqual(addit(-2), 3, "Testing the function addit with input -2 (should be 3)")
    self.assertEqual(addit(0), 5, "Testing the function addit with input 0 (should be 5)")
myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_functions-can-call-other-functions-composition.html#ac4_10_1" class="internal">in-context</a></span>
</body>
</html>
