<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.5.10</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_5_6">
<div class="ac_question col-md-12" id="foppff_1_ac4_5_6_question"><div class="para" id="p-1079">
<dfn class="terminology">10.</dfn> Write a function called <code class="code-inline tex2jax_ignore">subtract_three</code> that takes an integer or any number as input, and returns that number minus three.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_5_6_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(subtract_three(9), 6, "Testing the subtract_three function on input 9.")
        self.assertEqual(subtract_three(-5), -8, "Testing the subtract_three function on input -5.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_returning-a-value-from-a-function.html#ac4_5_6" class="internal">in-context</a></span>
</body>
</html>
