<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.5.12</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_5_8">
<div class="ac_question col-md-12" id="foppff_1_ac4_5_8_question"><div class="para" id="p-1081">
<dfn class="terminology">12.</dfn> Write a function named <code class="code-inline tex2jax_ignore">intro</code> that takes a string as input. This string is intended to be a person’s name and the output is a standardized greeting. For example, given the string “Becky” as input, the function should return: “Hello, my name is Becky and I love COMP 1701.”</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_5_8_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(intro("Mike"), "Hello, my name is Mike and I love COMP 1701.", "Testing the intro function on input 'Mike'.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_returning-a-value-from-a-function.html#ac4_5_8" class="internal">in-context</a></span>
</body>
</html>
