<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_15_2">
<div class="ac_question col-md-12" id="foppff_1_ac4_15_2_question"><div class="para" id="p-1380">Write a function called <code class="code-inline tex2jax_ignore">add</code> that takes any number as its input and returns that sum with 2 added.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_15_2_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testTwo(self):
    self.assertEqual(add(-2), 0, "Testing that add(-2) returns 0")
    self.assertEqual(add(6), 8, "Testing that add(6) returns 8")
    self.assertEqual(add(4), 6, "Testing that add(4) returns 6")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_chapter-assessment.html#ac4_15_2" class="internal">in-context</a></span>
</body>
</html>
