<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.5.13</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac4_5_9">
<div class="ac_question col-md-12" id="foppff_1_ac4_5_9_question"><div class="para" id="p-1082">
<dfn class="terminology">13.</dfn> Write a function called <code class="code-inline tex2jax_ignore">s_change</code> that takes one string as input and returns that string, concatenated with the string “for fun.”.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac4_5_9_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(s_change("Coding"), "Coding for fun." ,"Testing the function s_change with input coding")
        self.assertEqual(s_change("We go to the beach"), "We go to the beach for fun." , "Testing the function s_change with input We go to the beach")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="functions_returning-a-value-from-a-function.html#ac4_5_9" class="internal">in-context</a></span>
</body>
</html>
