<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">5.6.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac5_6_4">
<div class="ac_question col-md-12" id="foppff_1_ac5_6_4_question"><div class="para" id="p-1537">Create a variable, <code class="code-inline tex2jax_ignore">b</code>, and assign it the value of <code class="code-inline tex2jax_ignore">15</code>. Then, write code to see if the value <code class="code-inline tex2jax_ignore">b</code> is greater than that of <code class="code-inline tex2jax_ignore">a</code>. If it is, <code class="code-inline tex2jax_ignore">a</code>’s value should be multiplied by 2. If the value of <code class="code-inline tex2jax_ignore">b</code> is less than or equal to <code class="code-inline tex2jax_ignore">a</code>, nothing should happen. Finally, create variable <code class="code-inline tex2jax_ignore">c</code> and assign it the value of the sum of <code class="code-inline tex2jax_ignore">a</code> and <code class="code-inline tex2jax_ignore">
        b</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac5_6_4_editor" data-question_label="" data-codelens="true">a = 20

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testTwoA(self):
    self.assertEqual(a, 20, "Testing that a has the correct value.")

  def testTwoB(self):
    self.assertEqual(c, 35, "Testing that c has the correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="conditionals_conditional-execution-binary-selection.html#ac5_6_4" class="internal">in-context</a></span>
</body>
</html>
