<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">6.6.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac6_6_4">
<div class="ac_question col-md-12" id="foppff_1_ac6_6_4_question"><div class="para" id="p-1709">Write a function called <code class="code-inline tex2jax_ignore">stop_at_z</code> that iterates through a list of strings. Using a while loop, append each string to a new list until the string that appears is “z”. The function should return the new list.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac6_6_4_editor" data-question_label="" data-codelens="true">def stop_at_z():

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(stop_at_z(['c', 'b', 'd', 'zebra', 'h', 'r', 'z', 'm', 'a', 'k']), ['c', 'b', 'd', 'zebra', 'h', 'r'], "Testing the function stop_at_z on the input ['c', 'b', 'd', 'zebra', 'h', 'r', 'z', 'm', 'a', 'k'].")
      self.assertEqual(stop_at_z(['zoo', 'zika', 'ozzie', 'pizzazz', 'z', 'pizza', 'zap', 'haze']), ['zoo', 'zika', 'ozzie', 'pizzazz'], "Testing the function stop_at_z on the input ['zoo', 'zika', 'ozzie', 'pizzazz', 'z', 'pizza', 'zap', 'haze'].")
      self.assertEqual(stop_at_z(['z']), [], "Testing the function stop_at_z on the input ['z'].")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="iteration_chapter-assessment.html#ac6_6_4" class="internal">in-context</a></span>
</body>
</html>
