<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">6.2.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac6_2_2">
<div class="ac_question col-md-12" id="foppff_1_ac6_2_2_question"><div class="para" id="p-1667">Write a while loop that is initialized at 0 and stops at 15. If the counter is an even number, add double its value to <code class="code-inline tex2jax_ignore">total_score</code>. If the counter is an odd number, subtract half its value from <code class="code-inline tex2jax_ignore">total_score</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac6_2_2_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
     self.assertEqual(total_score, 80.0, "Testing that eve_nums has been assigned the correct elements")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="more-about-iteration_the-while-statement.html#ac6_2_2" class="internal">in-context</a></span>
</body>
</html>
