<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">6.6.5</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac6_6_5">
<div class="ac_question col-md-12" id="foppff_1_ac6_6_5_question"><div class="para" id="p-1710">Below is a for loop that works. Underneath the for loop, rewrite the problem so that it does the same thing, but using a while loop instead of a for loop. Assign the accumulated total in the while loop code to the variable <code class="code-inline tex2jax_ignore">sum2</code>. Once complete, sum2 should equal sum1.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac6_6_5_editor" data-question_label="" data-codelens="true">sum1 = 0

lst = [65, 78, 21, 33]

for x in lst:
    sum1 = sum1 + x

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFive(self):
      self.assertEqual(sum2, 197, "Testing that sum2 is assigned to correct value.")
      self.assertIn('while', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="iteration_chapter-assessment.html#ac6_6_5" class="internal">in-context</a></span>
</body>
</html>
