<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">7.5.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac7_5_5">
<div class="ac_question col-md-12" id="foppff_1_ac7_5_5_question"><div class="para" id="p-1776">Count the number of characters in string <code class="code-inline tex2jax_ignore">str1</code>. Do not use <code class="code-inline tex2jax_ignore">len()</code>. Save the number in variable <code class="code-inline tex2jax_ignore">numbs</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac7_5_5_editor" data-question_label="" data-codelens="true">str1 = "I like nonsense, it wakes up the brain cells. Fantasy is a necessary ingredient in living."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testEight(self):
      self.assertEqual(numbs, 90, "Testing that numbs is assigned to correct values.")
      self.assertNotIn("len(", self.getEditorText(), "Testing your code (Don't worry about actual and expected values)."))

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="more-about-iteration_the-accumulator-pattern.html#ac7_5_5" class="internal">in-context</a></span>
</body>
</html>
