<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h4 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.3.6</span><span class="period">.</span>
</h4>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac8_3_6">
<div class="ac_question col-md-12" id="foppff_1_ac8_3_6_question"><div class="para" id="p-1878">Assign the value of the last chacter of <code class="code-inline tex2jax_ignore">lst</code> to the variable <code class="code-inline tex2jax_ignore">output</code>. Do this so that the length of lst doesn’t matter.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac8_3_6_editor" data-question_label="" data-codelens="true">lst = "Every chess or checkers game begins from the same position and has a finite number of moves that can be played. While the number of possible scenarios and moves is quite large, it is still possible for computers to calculate that number and even be programmed to respond well against a human player..."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testThree(self):
      self.assertEqual(output, ".", "Testing that output value is assigned to correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="sequences_index-operator-working-with-the-characters-of-a-string.html#ac8_3_6" class="internal">in-context</a></span>
</body>
</html>
