<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.9.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac8_9_5">
<div class="ac_question col-md-12" id="foppff_1_ac8_9_5_question"><div class="para" id="p-2017">Create a variable <code class="code-inline tex2jax_ignore">output</code> and assign to it a list whose elements are the words in the string <code class="code-inline tex2jax_ignore">str1</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac8_9_5_editor" data-question_label="" data-codelens="true">str1 = "OH THE PLACES YOU'LL GO"

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testSix(self):
      self.assertEqual(output, ["OH", "THE", "PLACES", "YOU'LL", "GO"], "Testing that output value is assigned to correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="sequences_splitting-and-joining-strings.html#ac8_9_5" class="internal">in-context</a></span>
</body>
</html>
