<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.10.9</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac8_10_3">
<div class="ac_question col-md-12" id="foppff_1_ac8_10_3_question"><div class="para" id="p-2078">Write code to determine how many 9’s are in the list <code class="code-inline tex2jax_ignore">nums</code> and assign that value to the variable <code class="code-inline tex2jax_ignore">how_many</code>. Do not use a for loop to do this.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac8_10_3_editor" data-question_label="" data-codelens="true">nums = [4, 2, 23.4, 9, 545, 9, 1, 234.001, 5, 49, 8, 9 , 34, 52, 1, -2, 9.1, 4]


=====

from unittest.gui import TestCaseGui
import re

class myTests(TestCaseGui):

  def test_output(self):
    self.assertEqual(how_many, 3, "Testing that how_many is set correctly.")
    self.assertNotIn('for', self.getEditorText(), "Testing that you didn't use a for loop.")
    self.assertFalse(re.search(r'how_many\s*=\s*3', self.getEditorText()), "Hardcode check")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="sequences_chapter-assessment.html#ac8_10_3" class="internal">in-context</a></span>
</body>
</html>
