<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h4 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.3.5</span><span class="period">.</span>
</h4>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac8_3_5">
<div class="ac_question col-md-12" id="foppff_1_ac8_3_5_question"><div class="para" id="p-1877">Assign the value of the 23rd element of <code class="code-inline tex2jax_ignore">l</code> to the variable <code class="code-inline tex2jax_ignore">checking</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac8_3_5_editor" data-question_label="" data-codelens="true">l = ("hi", "goodbye", "python", "106", "506", 91, ['all', 'Paul', 'Jackie', "UMSI", 1, "Stephen", 4.5], 109, "chair", "pizza", "wolverine", 2017, 3.92, 1817, "account", "readings", "papers", 12, "facebook", "twitter", 193.2, "snapchat", "leaders and the best", "social", "1986", 9, 29, "holiday", ["women", "olympics", "gold", "rio", 21, "2016", "men"], "26trombones")

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(checking, "leaders and the best", "Testing that checking has the correct element assigned.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="sequences_index-operator-working-with-the-characters-of-a-string.html#ac8_3_5" class="internal">in-context</a></span>
</body>
</html>
