<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h4 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.3.4</span><span class="period">.</span>
</h4>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac8_3_4">
<div class="ac_question col-md-12" id="foppff_1_ac8_3_4_question"><div class="para" id="p-1876">Assign the value of the 34th element of <code class="code-inline tex2jax_ignore">lst</code> to the variable <code class="code-inline tex2jax_ignore">output</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac8_3_4_editor" data-question_label="" data-codelens="true">lst = ["hi", "morning", "dog", "506", "caterpillar", "balloons", 106, "yo-yo", "python", "moon", "water", "sleepy", "daffy", 45, "donald", "whiteboard", "glasses", "markers", "couches", "butterfly", "100", "magazine", "door", "picture", "window", ["Olympics", "handle"], "chair", "pages", "readings", "burger", "juggle", "craft", ["store", "poster", "board"], "laptop", "computer", "plates", "hotdog", "salad", "backpack", "zipper", "ring", "watch", "finger", "bags", "boxes", "pods", "peas", "apples", "horse", "guinea pig", "bowl", "EECS"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(output, "laptop", "Testing that output value is assigned to correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="sequences_index-operator-working-with-the-characters-of-a-string.html#ac8_3_4" class="internal">in-context</a></span>
</body>
</html>
