<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">8.10.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac8_10_2">
<div class="ac_question col-md-12" id="foppff_1_ac8_10_2_question"><div class="para" id="p-2032">Write code that combines the following variables so that the sentence “You are doing a great job, keep it up!” is assigned to the variable <code class="code-inline tex2jax_ignore">message</code>. Do not edit the values assigned to <code class="code-inline tex2jax_ignore">by</code>, <code class="code-inline tex2jax_ignore">az</code>, <code class="code-inline tex2jax_ignore">io</code>, or <code class="code-inline tex2jax_ignore">qy</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac8_10_2_editor" data-question_label="" data-codelens="true">by = "You are"
az = "doing a great "
io = "job"
qy = "keep it up!"


=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def test_output(self):
    self.assertEqual(by, 'You are', "Testing original variables.")
    self.assertEqual(az, 'doing a great ', "Testing original variables.")
    self.assertEqual(io, 'job', "Testing original variables.")
    self.assertEqual(qy, 'keep it up!', "Testing original variables.")
    self.assertEqual(message, 'You are doing a great job, keep it up!', "Testing that the value of message is what was expected.")
    self.assertNotIn("You are doing a great job, keep it up!", self.getEditorText(), "Testing for hardcoding (Don't worry about actual and expected values).")


myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="sequences_chapter-assessment.html#ac8_10_2" class="internal">in-context</a></span>
</body>
</html>
