<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">9.10.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac9_10_2">
<div class="ac_question col-md-12" id="foppff_1_ac9_10_2_question"><div class="para logical" id="p-2286"><ol class="decimal"><li id="li-281"><div class="para" id="p-2287">For each word in the list <code class="code-inline tex2jax_ignore">verbs</code>, add an -ing ending. Save this new list in a new list, <code class="code-inline tex2jax_ignore">ing</code>.</div></li></ol></div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac9_10_2_editor" data-question_label="" data-codelens="true">verbs = ["kayak", "cry", "walk", "eat", "drink", "fly"]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testSix(self):
      self.assertEqual(ing, ['kayaking', 'crying', 'walking', 'eating', 'drinking', 'flying'], "Testing that the variable ing has the correct value.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="transforming-sequences_the-accumulator-pattern-with-lists.html#ac9_10_2" class="internal">in-context</a></span>
</body>
</html>
