<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">9.11.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac9_11_6">
<div class="ac_question col-md-12" id="foppff_1_ac9_11_6_question"><div class="para" id="p-2304">Assign an empty string to the variable <code class="code-inline tex2jax_ignore">output</code>. Using the <code class="code-inline tex2jax_ignore">range</code> function, write code to make it so that the variable <code class="code-inline tex2jax_ignore">output</code> has 35 <code class="code-inline tex2jax_ignore">a</code> s inside it (like <code class="code-inline tex2jax_ignore">"aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa"</code>). Hint: use the accumulation pattern!</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac9_11_6_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testThree(self):
      self.assertEqual(output, "aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa", "Testing that output has the correct value.")
      self.assertNotIn("aaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaaa", self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="transforming-sequences_the-accumulator-pattern-with-strings.html#ac9_11_6" class="internal">in-context</a></span>
</body>
</html>
