<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">9.10.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac9_10_3">
<div class="ac_question col-md-12" id="foppff_1_ac9_10_3_question"><div class="para" id="p-2288">Given the list of numbers, <code class="code-inline tex2jax_ignore">numbs</code>, create a new list of those same numbers increased by 5. Save this new list to the variable <code class="code-inline tex2jax_ignore">newlist</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac9_10_3_editor" data-question_label="" data-codelens="true">numbs = [5, 10, 15, 20, 25]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFour(self):
      self.assertEqual(newlist, [10, 15, 20, 25, 30], "Testing that the newlist value contains the correct elements.")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="transforming-sequences_the-accumulator-pattern-with-lists.html#ac9_10_3" class="internal">in-context</a></span>
</body>
</html>
