<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">9.11.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac9_11_2">
<div class="ac_question col-md-12" id="foppff_1_ac9_11_2_question"><div class="para logical" id="p-2302"><ol class="decimal"><li id="li-282"><div class="para" id="p-2303">For each character in the string already saved in the variable <code class="code-inline tex2jax_ignore">str1</code>, add each character to a list called <code class="code-inline tex2jax_ignore">chars</code>.</div></li></ol></div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac9_11_2_editor" data-question_label="" data-codelens="true">str1 = "I love python"
# HINT: what's the accumulator? That should go here.

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testTwo(self):
      self.assertEqual(chars, ['I', ' ', 'l', 'o', 'v', 'e', ' ', 'p', 'y', 't', 'h', 'o', 'n'], "Testing that chars is assigned to correct values.")
      self.assertIn('append', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="transforming-sequences_the-accumulator-pattern-with-strings.html#ac9_11_2" class="internal">in-context</a></span>
</body>
</html>
