<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">9.10.5</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_ac9_10_4">
<div class="ac_question col-md-12" id="foppff_1_ac9_10_4_question"><div class="para" id="p-2289">Given the list of numbers, <code class="code-inline tex2jax_ignore">numbs</code>, modifiy the list <code class="code-inline tex2jax_ignore">numbs</code> so that each of the original numbers are increased by 5.  Note this is not an accumulator pattern problem, but its a good review.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_ac9_10_4_editor" data-question_label="" data-codelens="true">numbs = [5, 10, 15, 20, 25]

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testFive(self):
      self.assertEqual(numbs, [10, 15, 20, 25, 30], "Testing that numbs is assigned to correct values.")
      self.assertNotIn("append", self.getEditorText(), "Do not use append")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="transforming-sequences_the-accumulator-pattern-with-lists.html#ac9_10_4" class="internal">in-context</a></span>
</body>
</html>
