<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">2.17.1</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_assess_ps_01_01">
<div class="ac_question col-md-12" id="foppff_1_assess_ps_01_01_question"><div class="para" id="p-678">There is a function we are providing in for you in this problem called <code class="code-inline tex2jax_ignore">square</code>. It takes one integer and returns the square of that integer value. Write code to assign a variable called <code class="code-inline tex2jax_ignore">xyz</code> the value <code class="code-inline tex2jax_ignore">5*5</code> (five squared). Use the square function, rather than just multiplying with <code class="code-inline tex2jax_ignore">*</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_assess_ps_01_01_editor" data-question_label="" data-codelens="true">xyz = ""

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

def testOne(self):
self.assertEqual(type(xyz), type(3), "Checking type of xyz")
self.assertEqual(xyz, 25, "Checking if xyz is 25")
self.assertIn('square', self.getEditorText(), "Testing that 'square' is in your code. (Don't
worry about Actual and Expected Values.)")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="simple-python-data_chapter-assessment.html#assess_ps_01_01" class="internal">in-context</a></span>
</body>
</html>
