<!DOCTYPE html>
<html lang="en-US">
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<head>
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<meta name="robots" content="noindex, nofollow">
</head>
<body class="ignore-math">
<article class="exercise exercise-like"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">2.17.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="foppff_1_assess_ps_01_02">
<div class="ac_question col-md-12" id="foppff_1_assess_ps_01_02_question"><div class="para" id="p-679">Write code to assign the number of <em class="emphasis">characters</em> in the string <code class="code-inline tex2jax_ignore">rv</code> to a variable <code class="code-inline tex2jax_ignore">num_chars</code>.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="foppff_1_assess_ps_01_02_editor" data-question_label="" data-codelens="true">rv = """Once upon a midnight dreary, while I pondered, weak and weary,
Over many a quaint and curious volume of forgotten lore,
While I nodded, nearly napping, suddenly there came a tapping,
As of some one gently rapping, rapping at my chamber door.
'Tis some visitor, I muttered, tapping at my chamber door;
Only this and nothing more."""

# Write your code here!

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

def testOne(self):
self.assertEqual(num_chars, len(rv), "Testing that num_chars has been set to the length of
rv")

myTests().main()
</textarea>
</div></div></div></article><span class="incontext"><a href="simple-python-data_chapter-assessment.html#assess_ps_01_02" class="internal">in-context</a></span>
</body>
</html>
